<a href="https://colab.research.google.com/github/bragonzu/DATASCIENCE/blob/main/Introduction%20to%20Data%20Science%20in%20Python/Semana%202/Assignment2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**


Para esta tarea, verá los datos de 2017 sobre inmunizaciones de los CDC. Su archivo de datos para esta asignación se encuentra en assets / NISPUF17.csv. Una guía de usuarios de datos para esto, que necesitará para mapear las variables en los datos con las preguntas que se hacen, está disponible en assets / NIS-PUF17-DUG.pdf. Nota: es posible que deba ir a su árbol de Jupyter (haga clic en la imagen de Coursera) y navegar a la carpeta de activos de la asignación 2 para ver este archivo PDF).


## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 


##Pregunta 1
Escriba una función llamada "proporción_de_educación" que devuelva la proporción de niños en el conjunto de datos que tenían una madre con niveles de educación inferiores a la escuela secundaria (<12), la escuela secundaria (12), más que la escuela secundaria pero no un graduado universitario (> 12) y título universitario.

* Esta función debe devolver un diccionario en forma de (use los números correctos, no redondee los números): *
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    # your code goes here
    # YOUR CODE HERE
        import pandas as pd
        df= pd.read_csv("assets/NISPUF17.csv")
        n=df['EDUC1'].value_counts()/df['EDUC1'].shape[0]
        x=["less than high school", "high school", "more than high school but not college", "college"]
        di={}
        i=0
        for r in x:
            di[r]=n[i+1]
            i+=1
        
        return di
        
    #raise NotImplementedError()
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*

```
(2.5, 0.1)
```

Pregunta 2

Exploremos la relación entre recibir leche materna cuando era niño y recibir una vacuna contra la influenza estacional de un proveedor de atención médica. Devuelva una tupla del número promedio de vacunas contra la influenza para los niños que sabemos que recibieron leche materna cuando eran niños y los que saben que no.

Esta función debería devolver una tupla en el formulario (use los números correctos:

```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    # YOUR CODE HERE
    import pandas as pd
    import numpy as np
    df= pd.read_csv("assets/NISPUF17.csv") //lectura csv
    m1=df[df['CBF_01']==1]
    m2=df[df['CBF_01']==2]
    
    return np.mean(m1['P_NUMFLU']),np.mean(m2['P_NUMFLU'])
    #raise NotImplementedError()
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

---

Sería interesante ver si existe alguna evidencia de un vínculo entre la efectividad de la vacuna y el sexo del niño. Calcule la proporción de niños que contrajeron varicela pero fueron vacunados contra ella (al menos una dosis de varicela) versus aquellos que fueron vacunados pero no contrajeron varicela. Devuelve los resultados por sexo.

Esta función debería devolver un diccionario en forma de (use los números correctos):
''
     {"male": 0,2,
     "female": 0,4}
''

Nota: Para ayudar en la verificación, el valor de `chickenpox_by_sex () ['female']` que el autograder está buscando comienza con los dígitos `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    # YOUR CODE HERE
    import pandas as pd
    import numpy as np
    df= pd.read_csv("assets/NISPUF17.csv") //lectura csv
    am=df[(df['P_NUMVRC']>=1) & (df['HAD_CPOX']==1) & (df['SEX']==1)]
    af=df[(df['P_NUMVRC']>=1) & (df['HAD_CPOX']==1) & (df['SEX']==2)]
    bm=df[(df['P_NUMVRC']>=1) & (df['HAD_CPOX']==2) & (df['SEX']==1)]
    bf=df[(df['P_NUMVRC']>=1) & (df['HAD_CPOX']==2) & (df['SEX']==2)]
    
    rm=am.shape[0]/bm.shape[0]
    
    rf=af.shape[0]/bf.shape[0]
    r=[rm,rf]
    sex=['male','female']
    d={}
    i=0
    for s in sex:
        d[s]=r[i]
        i+=1
    return d
    #raise NotImplementedError()
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?


---

la correlación es una relación estadística entre dos variables. Si quisiéramos saber si las vacunas funcionan, podríamos ver la correlación entre el uso de la vacuna y si resulta en la prevención de la infección o enfermedad [1]. En esta pregunta, debe ver si existe una correlación entre haber tenido varicela y el número de dosis de vacuna contra la varicela administradas (varicela).

Algunas notas sobre la interpretación de la respuesta. Had_chickenpox_column es 1 (para sí) o 2 (para no), y num_chickenpox_vaccine_column es el número de dosis que se le ha dado a un niño de la vacuna contra la varicela. Una correlación positiva (por ejemplo, corr> 0) significa que un aumento en had_chickenpox_column (lo que significa más no) también aumentaría los valores de num_chickenpox_vaccine_column (lo que significa más dosis de vacuna). Si existe una correlación negativa (p. Ej., Corr <0), indica que haber tenido varicela está relacionado con un aumento en el número de dosis de vacuna.

Además, pval es la probabilidad de que observemos una correlación entre had_chickenpox_column y num_chickenpox_vaccine_column que es mayor o igual que un valor particular ocurrido por casualidad. Un pval pequeño significa que es muy poco probable que la correlación observada ocurra por casualidad. En este caso, pval debe ser muy pequeño (terminará en e-18 indicando un número muy pequeño).

[1] Esta no es realmente la imagen completa, ya que no estamos viendo cuándo se administró la dosis. Es posible que los niños tuvieran varicela y luego sus padres fueran a ponerles la vacuna. ¿Este conjunto de datos tiene los datos que necesitaríamos para investigar el momento de la dosis?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    
    df = pd.read_csv("assets/NISPUF17.csv")
    
    a = df[(df['P_NUMVRC'] >=0) & (df['HAD_CPOX'] <= 2)]
    
    no_yes = a['HAD_CPOX']
    
    est_vaccine = a['P_NUMVRC']
    
    
    
    # this is just an example dataframe
    df=pd.DataFrame({"had_chickenpox_column":no_yes,
                   "num_chickenpox_vaccine_column":est_vaccine})

    # here is some stub code to actually run the correlation
    corr, pval=stats.pearsonr(df["had_chickenpox_column"],df["num_chickenpox_vaccine_column"])
    
    # just return the correlation
    return corr

    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```